# D2 : Data Wrangling

In this week's workbook, you'll work through an example in this to gain practice on how to:
* read a DataFrame into Python
* wrangle the data

# Part I: Setup

Data wrangling often requires additional functionality outside what's included in Python by default. For this, we'll import other functions from helpful packages.

**Import the following packages using their common shortened name found in parentheses:**

* `numpy` (`np`)
* `pandas` (`pd`)

<!-- grade_id: cell-b689d21f629d9dda, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[1], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-2935019a026ca3a1, points: 0.2, grade: true, solution: false -->
```python
assert np
assert pd
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[2], line 1
----> 1 assert np
      2 assert pd

NameError: name 'np' is not defined
```

**Run the following cell code to make things throughout the rest of this workbook a little prettier.** (Note: You don't have to edit code here, but are free to and see what changes to be sure you understand each line.)

```python
# Configure libraries

# Don't display too many rows/cols of DataFrames
pd.options.display.max_rows = 7
pd.options.display.max_columns = 8
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[3], line 4
      1 # Configure libraries
      2 
      3 # Don't display too many rows/cols of DataFrames
----> 4 pd.options.display.max_rows = 7
      5 pd.options.display.max_columns = 8

NameError: name 'pd' is not defined
```

**Read the CSV file at the URL https://raw.githubusercontent.com/fivethirtyeight/data/master/steak-survey/steak-risk-survey.csv into Python in and assign it to the variable `survey`**.

<!-- grade_id: cell-65e31c6985ba9dd2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# take a look at the data
survey
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[4], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()
      4 # take a look at the data
      5 survey

NotImplementedError:
```

<!-- grade_id: cell-0fdae0bebda7c436, points: 0.2, grade: true, solution: false -->
```python
assert survey.shape == (551, 15)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[5], line 1
----> 1 assert survey.shape == (551, 15)

NameError: name 'survey' is not defined
```

<!-- grade_id: cell-caef12a25ae8fc28, points: 0.2, grade: true, solution: false -->
```python
assert survey.columns[0] == 'RespondentID'
assert survey.iloc[1,1] == 'Lottery B'
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[6], line 1
----> 1 assert survey.columns[0] == 'RespondentID'
      2 assert survey.iloc[1,1] == 'Lottery B'

NameError: name 'survey' is not defined
```

These data contain survey responses from Americans who responded to a SurveyMonkey Audience poll. These data were used in the [FiveThirtyEight](https://fivethirtyeight.com) article: *[How Americans Like Their Steak](https://fivethirtyeight.com/features/how-americans-like-their-steak/)*

# Part II: Wrangling

**Write a line of code to look at the first few rows of the DataFrame**

<!-- grade_id: cell-c7cd7d4cfbf98114, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[7], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

What do you notice about the first row of the dataframe? Notice that it's not actually an observation from a respondent? **Remove this row from the dataset. Assign this back to the variable `survey`. Print the first few rows again to make sure you've accomplished this.**

<!-- grade_id: cell-0624532e0d0d4ce1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[8], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-3969e31aae66b836, points: 0.2, grade: true, solution: false -->
```python
assert survey.shape == (550, 15)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[9], line 1
----> 1 assert survey.shape == (550, 15)

NameError: name 'survey' is not defined
```

<!-- grade_id: cell-d3cf9af60d6a942b, points: 0.2, grade: true, solution: false -->
```python
assert survey.iloc[0,1] == 'Lottery B'
assert survey.columns[2] == 'Do you ever smoke cigarettes?'
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[10], line 1
----> 1 assert survey.iloc[0,1] == 'Lottery B'
      2 assert survey.columns[2] == 'Do you ever smoke cigarettes?'

NameError: name 'survey' is not defined
```

Notice that there are a lot of different questions that were asked of respondents (columns) and 550 people who responded to the survey (rows). 

**Print a list of all the column names in this DataFrame.**

We'll only end up working with a subset of these.

<!-- grade_id: cell-c91f7a1c6729ad43, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[11], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

Now we have a sense of what information is included in the dataset. In the coming weeks, we'll answer the following questions:
1. Who cheats more on their significant other - males or females?
2. Are cigarette smokers less likely to skydive?
3. Do people in New England gamble more than other parts of the country?

To answer these we'll only need data from *some* of the columns in the dataset.

Let's drop the columns we don't need. **Drop the first two columns from the dataset. This should still be assigned to the variable `survey`.**

<!-- grade_id: cell-8a1106312081bab5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[12], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-89822f596640cd4e, points: 0.2, grade: true, solution: false -->
```python
assert survey.shape == (550,13)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[13], line 1
----> 1 assert survey.shape == (550,13)

NameError: name 'survey' is not defined
```

<!-- grade_id: cell-7951ab60702eea97, points: 0.2, grade: true, solution: false -->
```python
assert survey.columns[0] == 'Do you ever smoke cigarettes?'
# drop the columns?
assert 'RespondentID' not in survey.columns

assert survey.iloc[540, -1] == 'Pacific'
assert survey.iloc[340, 8] == 'Male'
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[14], line 1
----> 1 assert survey.columns[0] == 'Do you ever smoke cigarettes?'
      2 # drop the columns?
      3 assert 'RespondentID' not in survey.columns

NameError: name 'survey' is not defined
```

Now that we've got the columns we want, let's clean up those column names. **Rename the columns in `survey` so the appropriate columns have the following names:**

* smoking
* alcohol
* gambling
* skydiving
* speeding
* cheated
* steak
* steak_preference
* gender
* age
* income 
* education
* region

Note that a list has been provided here, with these names, so you don't have to type them all:

```python
['smoking','alcohol','gambling',
 'skydiving','speeding', 'cheated',
 'steak', 'steak_preference','gender',
 'age', 'income', 'education', 'region']
```

Output:
```
['smoking',
 'alcohol',
 'gambling',
 'skydiving',
 'speeding',
 'cheated',
 'steak',
 'steak_preference',
 'gender',
 'age',
 'income',
 'education',
 'region']
```

<!-- grade_id: cell-ff9ae6330821a062, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[16], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-ac6bcfa7d1ac939d, points: 0.2, grade: true, solution: false -->
```python
assert survey.shape == (550,13)
assert list(survey) == ['smoking',
                        'alcohol',
                        'gambling',
                        'skydiving',
                        'speeding',
                        'cheated',
                        'steak',
                        'steak_preference',
                        'gender',
                        'age',
                        'income',
                        'education',
                        'region']
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[17], line 1
----> 1 assert survey.shape == (550,13)
      2 assert list(survey) == ['smoking',
      3                         'alcohol',
      4                         'gambling',
   (...)     13                         'education',
     14                         'region']

NameError: name 'survey' is not defined
```

We're in pretty good shape now. **Print the first few rows of the `survey` DataFrame to see what the data look like at this point.**

```python
survey.head()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[18], line 1
----> 1 survey.head()

NameError: name 'survey' is not defined
```

# Part III: Cleaning

**Now that we've go the data we need, let's get a sense for how much missing data there is in this dataset by determining how many null-containing rows there are in `survey`. Assign this value to the variable `null_rows`.**

<!-- grade_id: cell-5bc6a10438afe331, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[19], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-e43974ae43ee28d8, points: 0.2, grade: true, solution: false -->
```python
assert null_rows == 217
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[20], line 1
----> 1 assert null_rows == 217

NameError: name 'null_rows' is not defined
```

Good to know that lots of people didn't answer every question. We'll keep that in mind as we work with this dataset.

Simply dropping missing observations is typically not good practice; however, in this case we'll drop observations that have missing data across the entire row, as these are individuals who didn't participate in the survey at all. **Remove rows where ALL the columns have missing data for that participant.**

<!-- grade_id: cell-b8417c0e50579e72, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[21], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-27d421b2b5aaac0f, points: 0.21, grade: true, solution: false -->
```python
assert survey.shape == (541, 13)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[22], line 1
----> 1 assert survey.shape == (541, 13)

NameError: name 'survey' is not defined
```

**Print the first few rows to remind yourself what the data look like at this point.**

<!-- grade_id: cell-e1148cbe6b8da30b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[23], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

Note that the first row no longer has all missing data here. We've got a dataset we can work with now!

Great work on this workbook - be sure to **submit on datahub**! We'll continue to work with this dataset in section to answer our questions of interest. You can work on other parts of the class or help a classmate work through this - we always understand things best once we've had to explain them to someone else.**