# D5: Inference

In this notebook, we'll use statistical approaches to attempt to determine what countries the best soccer (football, for the rest of the world) players come from. Some argue that Brazil produces the best footballers, while others argue it's got to be France. (And other have different opinions still.) We'll investigate this question using FIFA data from 2020.

Run the following cell. These are all you need for the assignment. Do not import additional packages.

```python
# Imports 
%matplotlib inline

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import math
import seaborn as sns
sns.set()
sns.set_context('talk')

import warnings
warnings.filterwarnings('ignore')
pd.set_option("display.max_columns", 104)
import patsy
import statsmodels.api as sm
import scipy.stats as stats
from scipy.stats import ttest_ind, chisquare, normaltest
# Note: the statsmodels import may print out a 'FutureWarning'. Thats fine.
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[1], line 15
     13 warnings.filterwarnings('ignore')
     14 pd.set_option("display.max_columns", 104)
---> 15 import patsy
     16 import statsmodels.api as sm
     17 import scipy.stats as stats

ModuleNotFoundError: No module named 'patsy'
```

## Part 1: Data & Wrangling

Fixing messy data makes up a large amount of the work of being a data scientist. 

The real world produces messy measurements and it is your job to find ways to standardize your data such that you can make useful analyses out of it. 

In this section, you will learn, and practice, how to successfully deal with unclean data.

### 1a) Load the data
Import datafile `players.csv` into a DataFrame called `df`.

Note these data come from [SOFIFA](https://sofifa.com/). Similar data have been scraped and shared on [Kaggle](https://www.kaggle.com/stefanoleone992/fifa-20-complete-player-dataset?select=players_20.csv); however the data we'll use here is distinct from that dataset.

<!-- grade_id: cell-252b5cfd5c8ceff2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[2], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: 1a_ans, points: 0.1, grade: true, solution: false -->
```python
assert isinstance(df, pd.DataFrame)
assert df.shape == (18278, 15)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[3], line 1
----> 1 assert isinstance(df, pd.DataFrame)
      2 assert df.shape == (18278, 15)

NameError: name 'df' is not defined
```

```python
# Check out the data
df.head(5)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[4], line 2
      1 # Check out the data
----> 2 df.head(5)

NameError: name 'df' is not defined
```

In the rest of Part 1, we will work on writing code, organized into functions that will allow us to transform similar respones into the same value. We will call this process: standardizing the data. 

The cell below provides an example for the kind of code you will need to write to answer this question. This example is separate from our actual data, and is a potential function we might use to standardize messy data - in this case, hypothetical data to the question 'What is your favourite major python version?'. 

Note some things used in this example that you need to use to standardize data:
- string methods, such as `lower` and `strip` to transform strings
- the `replace` string method, to replace a set of characters with something else
- if/else statements that check what's in our string (number, letters, etc)
- type casting, for example using `float()` to turn a variable into a float
- using `np.nan` (which stands for 'not a number') to denote missing or unknown data

The height column of the dataset is given and it is represented in either inches or centimeters. But, the data type is a string data type, the task here is to extract the number (in centimeters) and store it in the height column. The function given below `standardize_height` performs this task.

After a quick glance at the dataset, you would know that there are 4 possible ways the height column is filled as:
- It is represented in inches followed by `in` or `inches` to convey that the number is in inches.
- It is represented in centimeters followed by `cms` or `centimeters` to convey that the number is in centimeters.

The following is what the function does:
- Convert all the characters to lower case using `lower'
- Drop all white spaces using `strip`
- Check if it is inches, if so replace the keywords as empty strings and then convert the string datatype to float and the convert it to centimeters and finally return this float
- The same operation is as above is done, but the conversion from inches to centimeters isn't required here.

A thing to note here :
- The conditions for `inches` and `in` is checked simultaneously, though it could be done seperately (that is totally acceptable, this function is written to point out some common errors). Note that `replace` is used to replace `inches` first and then `in`, if it were the other way around, if the string contains `inches`, `in` would be replaced by an empty string and we would be left with `ches`, the conversion wouldn't be successful. 


```python
def standardize_height(str_in):
    
    try:
        str_in = str_in.lower()
        str_in = str_in.strip()

        if 'in' in str_in or 'inches' in str_in:
            str_in = str_in.replace('inches', '')
            str_in = str_in.replace('in', '')
            output = float(str_in)
            output = 2.54 * output
        elif 'cms' in str_in or 'centimeters' in str_in:
            str_in = str_in.replace('cms', '')
            str_in = str_in.replace('centimeters', '')
            output = float(str_in)
        else:
            output = np.nan
    except:
        output = np.nan
        
    return output
```

```python
assert standardize_height('180 cms') == 180
assert np.isclose(standardize_height('70 inches'), 177.8)
assert np.isnan(standardize_height('LOL'))
```

We then have to apply the transformation using the function we just defined.

```python
df['height'] = df['height'].apply(standardize_height)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[7], line 1
----> 1 df['height'] = df['height'].apply(standardize_height)

NameError: name 'df' is not defined
```

```python
# just run this cell
df.head(5)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[8], line 2
      1 # just run this cell
----> 2 df.head(5)

NameError: name 'df' is not defined
```

### 1b) Standardize 'weight' function

Next let's check the 'weight' column. 

Check the different responses received for weight, including how many of each response we have

```python
# run this to see different weight input data
df['weight'].value_counts()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[9], line 2
      1 # run this to see different weight input data
----> 2 df['weight'].value_counts()

NameError: name 'df' is not defined
```

Using a similar approach to what we used for 'height', you'll write a `standardize_weight` function.

To do this you'll:
- `try` to:
    - convert all text to lowercase
    - use the string method `strip()` to remove leading and trailing characters from the gender value
    - use an `if/elif/else` to:
        - check if 'pounds' or 'lbs' exists. If so replace these words with an empty string, convert it to float, then to kilograms, and return it.
        - check if 'kgs' or 'kilograms' exists. If so replace these words with an empty string and convert it to float and return this value.
        - return `np.nan` otherwise.
- `except`: return `np.nan`
    
Note: use `value in kilograms = 0.453592 * value in pounds`.

<!-- grade_id: cell-2720852762f49edd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[10], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
standardize_weight('200 lbs')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[11], line 1
----> 1 standardize_weight('200 lbs')

NameError: name 'standardize_weight' is not defined
```

<!-- grade_id: cell-a3748b058a2402f3, points: 0.5, grade: true, solution: false -->
```python
assert np.isclose(standardize_weight('200 Pounds'), 90.7, 0.1)
assert standardize_weight('80.0 kgs') == 80
assert np.isclose(standardize_weight('200 lbs'), 90.7, 0.1)
assert standardize_weight('80.0 KILOGRAMS') == 80
assert np.isnan(standardize_weight('unicorn'))
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[12], line 1
----> 1 assert np.isclose(standardize_weight('200 Pounds'), 90.7, 0.1)
      2 assert standardize_weight('80.0 kgs') == 80
      3 assert np.isclose(standardize_weight('200 lbs'), 90.7, 0.1)

NameError: name 'standardize_weight' is not defined
```

### 1c) Transform 'weight' column

Apply the transformation, meaning, use your function and standardize weight in `df`

<!-- grade_id: cell-291cf93820c28b6d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[13], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
# Check the results
df['weight'].unique()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[14], line 2
      1 # Check the results
----> 2 df['weight'].unique()

NameError: name 'df' is not defined
```

<!-- grade_id: cell-2c24bd8912a04a3f, points: 0.25, grade: true, solution: false -->
```python
assert len(df['weight']) == 18278
assert sum(df['weight'].isnull()) == 5
assert df['weight'].median() == 75.
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[15], line 1
----> 1 assert len(df['weight']) == 18278
      2 assert sum(df['weight'].isnull()) == 5
      3 assert df['weight'].median() == 75.

NameError: name 'df' is not defined
```

```python
df.head(5)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[16], line 1
----> 1 df.head(5)

NameError: name 'df' is not defined
```

## Part 2: EDA

Now that our data are ready to go, let's understand a bit more about our data...

### 2a) Nationality

We're ultimately going to be comparing French and Brazilian Players...so we should probably get a sense of how many different nationalities are represented in our dataset, by running the code below: 

```python
# determine how many different nationalities
df['nationality'].nunique()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[17], line 2
      1 # determine how many different nationalities
----> 2 df['nationality'].nunique()

NameError: name 'df' is not defined
```

We're most interested in the countries that produce the most/best players, so let's generate a barplot of the nations that proudce the top 10 most players:

<!-- grade_id: cell-f38279d756a7a299, points: 0.2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[18], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

From the plot you generated, you should see that England is the most common nationality, but that France and Brazil (our countries of interest) are in the top 10.

### 2b) Scatterplot Matrix

Using `scatter_matrix`, from `pandas`' `plotting` module, plot `df` for the 5 columns `age`, `height`, `weight`, `'potential`, and `overall`. Assign it to a variable called `fig`. 

<!-- grade_id: 2a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[19], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: 2a_ans, points: 0.2, grade: true, solution: false -->
```python
assert np.all(fig)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[20], line 1
----> 1 assert np.all(fig)

NameError: name 'fig' is not defined
```

From this plot, you should get a sense of the range/possible values for each variable and see that potential and overall are certainly related to one another, as are height and weight. However, as we would expect, age is not closely related to any of the other variables.

## Part 3: Analysis

While we haven't fully explored the data, for the purposes of this workbook, let's get to answering some questions! For this, we're only going to focus in on players from two countries: France and Brazil and we'll answer the following questions:

1. Do the heights of players differ between Brazil and France?
2. Do the weights of players differ between Brazil and France?
4. What is the effect of nationality on overall rating?

### 3a) Subset Data

Obtain a subset of the original dataframe `df`, including only players whose `nationality` is either 'Brazil' or 'France'. Store this in `df_sub`

<!-- grade_id: cell-96f1c4712855f7e8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[21], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-f615cf4869353bd1, points: 0.1, grade: true, solution: false -->
```python
assert df_sub.shape == (1808, 15)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[22], line 1
----> 1 assert df_sub.shape == (1808, 15)

NameError: name 'df_sub' is not defined
```

### 3b) Weight

Extract the weight data for all French players, storing it in `w_france`.

Extract the weight data for all Brazilian players, storing it in `w_brazil`.

```python
w_france = df[df['nationality'] == 'France']['weight']
w_brazil = df[df['nationality'] == 'Brazil']['weight']
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[23], line 1
----> 1 w_france = df[df['nationality'] == 'France']['weight']
      2 w_brazil = df[df['nationality'] == 'Brazil']['weight']

NameError: name 'df' is not defined
```

### 3c) Difference in weight?

Carry out a t-test (using the `ttest_ind()` function) to compare the two distributions. Store the test statstic in the variable `t_val` and the p-value in the variable `p_val`.

<!-- grade_id: cell-8a7bb7133fdd9d77, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[24], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

```python
# look at results
t_val, p_val
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[25], line 2
      1 # look at results
----> 2 t_val, p_val

NameError: name 't_val' is not defined
```

<!-- grade_id: cell-5cd84d5270b2fb06, points: 0.2, grade: true, solution: false -->
```python
assert np.isclose(t_val, -1.76, atol = 0.01)
assert np.isclose(p_val, 0.079, atol = 0.01)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[26], line 1
----> 1 assert np.isclose(t_val, -1.76, atol = 0.01)
      2 assert np.isclose(p_val, 0.079, atol = 0.01)

NameError: name 't_val' is not defined
```

```python
# Check if statistical test passes significance
# using an alpha value of 0.01. This code provided.
if p_val < 0.01:
    print('There is a significant difference!')
else:
    print('There is NOT a significant difference!')
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[27], line 3
      1 # Check if statistical test passes significance
      2 # using an alpha value of 0.01. This code provided.
----> 3 if p_val < 0.01:
      4     print('There is a significant difference!')
      5 else:

NameError: name 'p_val' is not defined
```

### 3d) Height

Carry out the same process as above for height to determine if the heights of players differ between France and Brazil. Again, store the output in `t_val` and `p_val`

<!-- grade_id: cell-1d621c2b31180085, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[28], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-4f324c0577de3b9c, points: 0.15, grade: true, solution: false -->
```python
assert np.isclose(t_val, 2.19, atol = 0.01)
assert np.isclose(p_val, 0.028, atol = 0.01)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[29], line 1
----> 1 assert np.isclose(t_val, 2.19, atol = 0.01)
      2 assert np.isclose(p_val, 0.028, atol = 0.01)

NameError: name 't_val' is not defined
```

### 3e) Who's taller?

In 3d, you should have determined that there is a height difference between French and Brazilian soccer players...but who's taller? Determine the difference in means here, storing the mean height for each country in `mean_france` and `mean_brazil`, respectively.

<!-- grade_id: cell-0f20ac538b6415d7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[30], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-e6605462e4d4b155, points: 0.2, grade: true, solution: false -->
```python
assert np.isclose(mean_france, 181.34, atol = 0.01)
assert np.isclose(mean_brazil, 180.67, atol = 0.01)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[31], line 1
----> 1 assert np.isclose(mean_france, 181.34, atol = 0.01)
      2 assert np.isclose(mean_brazil, 180.67, atol = 0.01)

NameError: name 'mean_france' is not defined
```

### 3f) Overall

While height and weight may affect your soccer prowess, how you perform is what ultimately matters. Let's determine the relationship (French vs. Brazil) between nationality and `overall` score in `df_sub` using `OLS` from `statsmodels`.

The null hypothesis here would be that there is no relationship between nationality and `overall` score. 

Store the results from your analysis in the variable `res_1`.

<!-- grade_id: cell-adf747720a71f6a9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[32], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-efeecd54137aed39, points: 0.1, grade: true, solution: false -->
```python
assert isinstance(res_1, sm.regression.linear_model.RegressionResultsWrapper)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[33], line 1
----> 1 assert isinstance(res_1, sm.regression.linear_model.RegressionResultsWrapper)

NameError: name 'res_1' is not defined
```

```python
# look at results
print(res_1.summary())
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[34], line 2
      1 # look at results
----> 2 print(res_1.summary())

NameError: name 'res_1' is not defined
```

In this model, you should see that, on average, French players tend to have significantly lower overall scores, relative to Brazilians. Our first piece of evidence suggesting that Brazil may be the country that produces the best players! However, we only looked at two countries here...and only looked at one feature. We'd have to do a lot more digging/analysis before determining that conclusively!