# Case Study : Inference

The idea here is that we spent in class talking about inference and walking through some examples, and now it's your turn to do a (not quite-as-guided) inferential analysis!

Your goal is to, without much structure provided, consider what we've discussed in the course up to this point. Then, utilize the data provided to answer the question: *Do Pulitzers help newspapers keep readers?*, or more specifically *By looking at Pulitzer prizes awarded and changes in readership, can we determine what the effect of prestige is on the viewership at the 50 most popular newspapers between 2004 and 2013?*

The URL to the data for use in answering this question from FiveThirtyEight is https://raw.githubusercontent.com/fivethirtyeight/data/master/pulitzer/pulitzer-circulation-data.csv. 


# Part I : Setup & Data Wrangling

In this section you'll want to:
- import any packages you'll need for your analysis
    - a number are provided for you, but feel free to edit if there are other packages you want to use or changes to the input you want to make (Note: this is allowed in labs; in assignments you're limited to the imports we give you)
- read the dataset in
- understand what's going on in the data
    - how many observations?
    - what variables do you have? what variable types are here
    - which will you need to answer the question?

```python
# Import seaborn and apply its plotting styles
import seaborn as sns
sns.set(font_scale=2, style="white")

# import matplotlib
import matplotlib as mpl
import matplotlib.pyplot as plt
import matplotlib.style as style
# set plotting size parameter
plt.rcParams['figure.figsize'] = (17, 7)

# import pandas & numpy library
import pandas as pd
import numpy as np

# Statmodels & patsy
import patsy
import statsmodels.api as sm
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[1], line 17
     14 import numpy as np
     16 # Statmodels & patsy
---> 17 import patsy
     18 import statsmodels.api as sm

ModuleNotFoundError: No module named 'patsy'
```

**Store the dataset (see URL above) in a variable called `pulitzer`** and take a look at the data.

<!-- grade_id: cell-1bbc77943e957b82, points: 0, grade: false, solution: true -->
```python
# Get CSV directly from URL above
# YOUR CODE HERE
url = 'https://raw.githubusercontent.com/fivethirtyeight/data/master/pulitzer/pulitzer-circulation-data.csv'
pulitzer = pd.read_csv(url)
```

<!-- grade_id: cell-74d663e3c72968c5, points: 0.25, grade: true, solution: false -->
```python
# get shape of dataset
assert pulitzer.shape == (50, 7)
```

```python
# take a look at the data
pulitzer.describe()
```

Output:
```
Pulitzer Prize Winners and Finalists, 1990-2003  \
count                                        50.000000   
mean                                          8.340000   
std                                          12.886649   
min                                           0.000000   
25%                                           1.000000   
50%                                           4.000000   
75%                                           9.750000   
max                                          55.000000   

       Pulitzer Prize Winners and Finalists, 2004-2014  \
count                                        50.000000   
mean                                          6.720000   
std                                          12.142084   
min                                           0.000000   
25%                                           1.000000   
50%                                           3.000000   
75%                                           6.750000   
max                 
```

Okay, so we can see the data here. Each row (observation) contains data about a different newspaper. It looks like we have information from 50 different newspapers. We ultimately want to look at the relationship between number of pulitzers won and readership. We'd like to see this over time, but we may be restricted to just comparing early 2000s with 2013. This is a limitation of the analysis, but a good place to start. Now, it's time to wrangle...

```python
# why aren't daily circulation numbers in there?
pulitzer.dtypes
```

Output:
```
Newspaper                                            str
Daily Circulation, 2004                              str
Daily Circulation, 2013                              str
Change in Daily Circulation, 2004-2013               str
Pulitzer Prize Winners and Finalists, 1990-2003    int64
Pulitzer Prize Winners and Finalists, 2004-2014    int64
Pulitzer Prize Winners and Finalists, 1990-2014    int64
dtype: object
```

Ah, so circulation is an object....we want this to be numeric. Let's take care of that! **Convert the three Circulation columns to floats below.** Overwrite the data stored in those three columns (rather than creating any new columns).

<!-- grade_id: cell-59123be261bd76c6, points: 0, grade: false, solution: true -->
```python
# get rid of the commas in the numbers for Daily Circulation
# YOUR CODE HERE
pulitzer['Daily Circulation, 2004'] = pulitzer['Daily Circulation, 2004'].str.replace(',', '').astype(float)
pulitzer['Daily Circulation, 2013'] = pulitzer['Daily Circulation, 2013'].str.replace(',', '').astype(float)
pulitzer['Change in Daily Circulation, 2004-2013'] = pulitzer['Change in Daily Circulation, 2004-2013'].str.replace(',', '').str.replace('%', '').astype(float)

```

```python
# verify that that worked
pulitzer.dtypes
```

Output:
```
Newspaper                                              str
Daily Circulation, 2004                            float64
Daily Circulation, 2013                            float64
Change in Daily Circulation, 2004-2013             float64
Pulitzer Prize Winners and Finalists, 1990-2003      int64
Pulitzer Prize Winners and Finalists, 2004-2014      int64
Pulitzer Prize Winners and Finalists, 1990-2014      int64
dtype: object
```

<!-- grade_id: cell-111d2320b442873a, points: 0.25, grade: true, solution: false -->
```python
assert(pulitzer['Daily Circulation, 2004'].dtype == float)
assert(pulitzer['Daily Circulation, 2013'].dtype == float)
assert(pulitzer['Change in Daily Circulation, 2004-2013'].dtype == float)
```

# Part II : EDA & Visualization

Once you understand your dataset, you'll want to:
- understand the relationships and trends for the variables in your dataset
- generate exploratory visualizations


<!-- grade_id: cell-534db5706f829313, points: 0, grade: false, solution: true -->
```python
# take a look at some descriptive statistics
# YOUR CODE HERE

```

Here we see that the mean daily circulation decreased from 2004 to 2013. This is in line with what we may expect, given reports that newspaper readership is decreasing in the age of the Internet. But, we also see that there is a fair amount of variability. This also makes sense, as some newspapers are much more popular than others. 

The change in daily circulation is important as well and we see that this number is negative, thus capturing the overall downward trend we expected in readership over this time period.

When we start to consider Pulitzter Prize winners we see that the mean number has varied from 8.3 in 1990-2003 to 6.7 in 2004-2014. We note that the final column combines these data across the time period of 1990-2014. 

Let's generate some plots to understand the overall distribution. Generate the plot you think best first the comments provided below

<!-- grade_id: cell-66ea5a82c36a7937, points: 0, grade: false, solution: true -->
```python
## look at daily circulation distribution in 2004 and in 2013
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[10], line 3
      1 ## look at daily circulation distribution in 2004 and in 2013
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

You should see that most newspapers have smaller readership bases, with a few outlier newspapers. Now, let's look at the distribution of Pulitzter prize winners for the same time period.

<!-- grade_id: cell-16c19faf220f24e4, points: 0, grade: false, solution: true -->
```python
## look at pulitzer prize winner distributions
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[11], line 3
      1 ## look at pulitzer prize winner distributions
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

Here, you should see that over the two time periods of interest the distributions look similar, although we note that we are only looking at the histogram for 50 newspapers, so our observations are limited. Nevertheless, most of the data clusters around 0, with ony a few newspapers receiving more than 20 pulitzers over the time periods we're interested in. These data are clearly not Normally distributed, but are skewed right...something we'll want to keep in mind when we analyze these data.

<!-- grade_id: cell-a848ec43fe9792a3, points: 0, grade: false, solution: true -->
```python
# plot the distribution of *change* in daily circulation
# YOUR CODE HERE
raise NotImplementedError()

```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[12], line 3
      1 # plot the distribution of *change* in daily circulation
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

Here we see our final univariate plot of interest. You should now see that while there is an overall downward trend in daily circulation (mean ~ -25), there are papers whose readership has increased. That outlier to the right, let's check out which paper that is...

<!-- grade_id: cell-b604452cc25caea8, points: 0, grade: false, solution: true -->
```python
# determine which paper has a positive change in circulation much larger than other papers
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[13], line 3
      1 # determine which paper has a positive change in circulation much larger than other papers
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

The paper you see above should make enough sense with your intution about popular publications. 

We also note that two newspapers have a -100 change in daily circulation. This is an interesting data point and something we want to investigate further. What's going on here? **Determine below which papers have -100 change in circulation.**

<!-- grade_id: cell-01cd489a7fd41f9f, points: 0, grade: false, solution: true -->
```python
# determine really large daily negative change in circulation
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[14], line 3
      1 # determine really large daily negative change in circulation
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

If we Google these papers we see that [Rocky Mountain News closed its doors in 2009](https://en.wikipedia.org/wiki/Rocky_Mountain_News). It makes sense that the readership decreased by -100 there. Now, what about New Orleans Times-Picayune? They [still have a website](https://www.nola.com/), so closure is not the same case here. When we dig a little deeper we see that [it is no longer published daily](https://en.wikipedia.org/wiki/The_Times-Picayune), explaining daily circulation having a measure of -100. We'll have to consider how to handle these data points. 

With that understanding of the readership numbers in our dataset, we'll continue our EDA to get a more specific understanding of the Pulitzer portion of the data. We now want to answer these questions before moving forward: 

1. What's a typical number of Pulitzers? What's that distribution look like?
1. Have the same newspapers consistently won more Pulitzers? 
2. Who has won the most Pulitzers? 
3. Anything look off?

**Generate the plots specified below, inspecting/understanding the output in each case.**

<!-- grade_id: cell-871c7ad90de9778f, points: 0, grade: false, solution: true -->
```python
## look at pulitzer prize winner distributions
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[15], line 3
      1 ## look at pulitzer prize winner distributions
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

Here you should see a distribution skewed right, with most publications receiving few awards and a few publications getting many more.

<!-- grade_id: cell-4a001e9a81ba4eec, points: 0, grade: false, solution: true -->
```python
# determine which publications won a whole bunch of Pulitzers (>80)
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[16], line 3
      1 # determine which publications won a whole bunch of Pulitzers (>80)
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

Above you should see here that these three publications earned the most overall, and the most in each of the time periods specified. They've been around a while and have been the most successful at winning Pulitzers. Let's see that trend overall across all papers in our dataset.

**Look at the trend of Pulitzer Prize winners 1990-2003 compared with 2004-2014 to determine if papers have similar trend over time.**

<!-- grade_id: cell-1b508a07589edfcc, points: 0, grade: false, solution: true -->
```python
# use fivethirtyeight style
# mostly for fun ; data are from them
style.use('fivethirtyeight') 

# plot relationship between pulitzer prize winners/finalists in each time period
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[17], line 7
      3 style.use('fivethirtyeight') 
      5 # plot relationship between pulitzer prize winners/finalists in each time period
      6 # YOUR CODE HERE
----> 7 raise NotImplementedError()

NotImplementedError:
```

Ok, so generally here you should see the same papers winning the most Pulitzers in the 90's are the same as those winning those in the 2000s.

<!-- grade_id: cell-71a02d011475713b, points: 0, grade: false, solution: true -->
```python
# who has won the most pulitzers during the years we're looking at? 
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[18], line 3
      1 # who has won the most pulitzers during the years we're looking at? 
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

As we saw above, these papers should like come to mind as some of the most popular papers nationally, so there are no huge surprises here. The number of Puitzers is relatively consistent across the two time periods.

Before we carry out our analysis, we're going to **remove the papers no longer in daily circulation from our analysis, as their data won't accurately reflect the relationship. Store this output in `pulitzer` again.**

<!-- grade_id: cell-b47b98ab3b40e3bd, points: 0, grade: false, solution: true -->
```python
# filter out two non daily-circulation papers.
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[19], line 3
      1 # filter out two non daily-circulation papers.
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-69a1b81fa0bb1c07, points: 0.25, grade: true, solution: false -->
```python
assert pulitzer.shape == (48, 7)
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[20], line 1
----> 1 assert pulitzer.shape == (48, 7)

AssertionError:
```

**Let's take a look at the change in daily cirulation after removing these outliers.**

<!-- grade_id: cell-16fb6a241f3d8536, points: 0, grade: false, solution: true -->
```python
# plot the distribution of daily change in circulation after outlier removal
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[21], line 3
      1 # plot the distribution of daily change in circulation after outlier removal
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

Now, after removing those two papers, if you look at daily circulation, the data are closer to normally distributed, but we note there are still papers that are outliers at the high end of circulation papers. 

# Part III : Inferential Analysis

Finally, answer the question! Carry out the inferential analysis required to answer your question of interest!

**So, you could just plot the relationship between the Pulittzers and daily circulation across the years in this dataset to get a sense of what's going on.**

This will help us start to answer: Does winning more Pulitzers result in your readership changing less?

<!-- grade_id: cell-b4cd1dfc53041ef2, points: 0, grade: false, solution: true -->
```python
# Relationship between the total number of Pulitzers and change in readership (daily circulation)
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[22], line 3
      1 # Relationship between the total number of Pulitzers and change in readership (daily circulation)
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

What you should see here is that there appears to be some general positive relationship but note that there are outliers in circulation, where circulation is much higher than number of Pulitzers won as well as outliers in Pulitzers won. Overall, most of the data are clustered around zero Pulitzers.  This is because the distributions of both of our variables are skewed right. This suggests that we may need to transform these data to use linear regression to ensure that the large outlier values are not driving our relationship.

So, as discussed in class: **Apply a log10-transformation the Pulitzer count data, with an offset of 0.1** This will shift the values away from being centered near zero when put on the log scale. Store this in a new column `pulitzer_log10`.

<!-- grade_id: cell-38d6e552b7f618e9, points: 0, grade: false, solution: true -->
```python
# log transform with a 0.1 offset

# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[23], line 4
      1 # log transform with a 0.1 offset
      2 
      3 # YOUR CODE HERE
----> 4 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-68b0ddccb8991617, points: 0.25, grade: true, solution: false -->
```python
assert pulitzer['pulitzer_log10'].shape == (48,)
assert pulitzer['pulitzer_log10'].min() == -1.0
assert np.isclose(pulitzer['pulitzer_log10'].max(), 2.07, 0.02)
```

Output:
```
File ~\AppData\Local\Programs\Python\Python312\Lib\site-packages\pandas\core\frame.py:4378, in DataFrame.__getitem__(self, key)
   4376 if self.columns.nlevels > 1:
   4377     return self._getitem_multilevel(key)
-> 4378 indexer = self.columns.get_loc(key)
   4379 if is_integer(indexer):
   4380     indexer = [indexer]

File ~\AppData\Local\Programs\Python\Python312\Lib\site-packages\pandas\core\indexes\base.py:3648, in Index.get_loc(self, key)
   3643     if isinstance(casted_key, slice) or (
   3644         isinstance(casted_key, abc.Iterable)
   3645         and any(isinstance(x, slice) for x in casted_key)
   3646     ):
   3647         raise InvalidIndexError(key) from err
-> 3648     raise KeyError(key) from err
   3649 except TypeError:
   3650     # If we have a listlike key, _check_indexing_error will raise
   3651     #  InvalidIndexError. Otherwise we fall through and re-raise
   3652     #  the TypeError.
   3653     self._check_indexing_error(key)

KeyError: 'pulitzer_log
```

<!-- grade_id: cell-69d66c93d2c7d07f, points: 0, grade: false, solution: true -->
```python
# visualize the distribtuion of the log10 column
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[25], line 3
      1 # visualize the distribtuion of the log10 column
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

Well, the distribution should appear less skewed...but what's going on down around -1? Well, this is because there were zeroes in the original dataset (see below). Due to this we used an offset of 0.1 in the log transformation above. These zeroes all show up at -1 as the $log_{10}(0+0.1) = -1$. 

```python
# this should show you all of the papers with zero pulitzer prizes
pulitzer[pulitzer['pulitzer_log10'] < 0]

```

Output:
```
File ~\AppData\Local\Programs\Python\Python312\Lib\site-packages\pandas\core\frame.py:4378, in DataFrame.__getitem__(self, key)
   4376 if self.columns.nlevels > 1:
   4377     return self._getitem_multilevel(key)
-> 4378 indexer = self.columns.get_loc(key)
   4379 if is_integer(indexer):
   4380     indexer = [indexer]

File ~\AppData\Local\Programs\Python\Python312\Lib\site-packages\pandas\core\indexes\base.py:3648, in Index.get_loc(self, key)
   3643     if isinstance(casted_key, slice) or (
   3644         isinstance(casted_key, abc.Iterable)
   3645         and any(isinstance(x, slice) for x in casted_key)
   3646     ):
   3647         raise InvalidIndexError(key) from err
-> 3648     raise KeyError(key) from err
   3649 except TypeError:
   3650     # If we have a listlike key, _check_indexing_error will raise
   3651     #  InvalidIndexError. Otherwise we fall through and re-raise
   3652     #  the TypeError.
   3653     self._check_indexing_error(key)

KeyError: 'pulitzer_log
```

Now it's time to do our analysis!

First, **plot the relationship between our two variables of interest** - be sure to consider the assumptions of linear regression discussed in class and determine if linear regression is appropriate.

Then, if appropriate, **use linear regression to determine the effect of log 10 pulitzer on change in daily circulation and take a look at the results**

<!-- grade_id: cell-5667dcd8e74452da, points: 0, grade: false, solution: true -->
```python
# plot relationship between the total number of Pulitzers and change readership
# YOUR CODE HERE
raise NotImplementedError()

f1 = plt.gcf()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[27], line 3
      1 # plot relationship between the total number of Pulitzers and change readership
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()
      5 f1 = plt.gcf()

NotImplementedError:
```

<!-- grade_id: cell-bc8f9229e99d7c49, points: 0.25, grade: true, solution: false -->
```python
assert f1.gca().has_data()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[28], line 1
----> 1 assert f1.gca().has_data()

NameError: name 'f1' is not defined
```

<!-- grade_id: cell-2eb058cc973e15cf, points: 0, grade: false, solution: true -->
```python
# carry out linear regression
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[29], line 3
      1 # carry out linear regression
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-85c9d5096faba7a3, points: 0.25, grade: true, solution: false -->
```python
# check that the correct model was used
assert 'OLS' in str(res_log.model)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[30], line 2
      1 # check that the correct model was used
----> 2 assert 'OLS' in str(res_log.model)

NameError: name 'res_log' is not defined
```

<!-- grade_id: cell-e98e4df45afe0530, points: 0.5, grade: true, solution: false -->
```python
# check that the model outputs are as expected
assert np.isclose(res_log.params[0], -32.5, 0.02)
assert np.isclose(res_log.params[1], 8.83, 0.02)
assert np.isclose(res_log.pvalues[1], 0.041, 0.02)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[31], line 2
      1 # check that the model outputs are as expected
----> 2 assert np.isclose(res_log.params[0], -32.5, 0.02)
      3 assert np.isclose(res_log.params[1], 8.83, 0.02)
      4 assert np.isclose(res_log.pvalues[1], 0.041, 0.02)

NameError: name 'res_log' is not defined
```

Here, we can't directly interpret our $\beta_1$ estimate (`coef`) on our original scale above because we log-transformed our data. In order to interpet 8.83 (what your $\beta_1$ should be), we would say that for every one unit increase in the number of pulitzers a publication has, the viewership increases by `~(coef/100)`, which here would be (8.83/100)...or 0.0883. That effect is pretty small. This is further reflected in the p-value (0.04), which suggests that we would see these results, or something more extreme due to random chance alone approximately 4% of the time.

We can therefore conclude that based on the data we have over these two time periods, there is a significant relationship between the number of Pulitzers a publication wins and the papers readership but that effect is quite minimal. We also know that this analysis is likely failing to take into account the effect of confounders we discussed in class, limiting our confidence in this result.

Note, of course, there are limitations to this analysis. There are other measures of prestige _and_ other measures of readership (electronic-only, for instance), that could help to answer this question. All analyses have limitations. Here, we've just noted a few. Be sure to consider all limitations for your final projects and include them in your Jupyter Notebook reports.